### CS102/CS103

Prof. Götz Pfeiffer<br />
School of Mathematics, Statistics and Applied Mathematics<br />
NUI Galway

# Computer Lab 8

## True or False

Each of the following statements is either true or false.
Delete the current content of the cell below the statement.
Then type in your answer as either `True` or `False`
(with a capital `T` or `F`).

* New objects are created by invoking a constructor.

<!-- grade_id: cell-32251635ac0af5f8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-c61d3857788f0519, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* Functions that live in objects are called instance variables.

<!-- grade_id: cell-937069a17599d431, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-94f5d5fdd250cf1b, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* The first parameter of a `python` method definition is called `this`.

<!-- grade_id: cell-26813d01043eb971, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-b497c88119322f1a, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* An object may have only one instance variable.

<!-- grade_id: cell-c6d5bb4df293fba7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-f9683c1579077ec3, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* In a `python` class, the constructor is called `__init__()`.

<!-- grade_id: cell-f994f82d6108817e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-0b361f608f71d179, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* A Docstring is the same thing as a comment.

<!-- grade_id: cell-c9c0fc7fb71c3728, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-b9ac2a319c004cb4, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* Method names should always begin with one or two underscores.

<!-- grade_id: cell-4b03f7f4b1b22f6a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-7497743302b7858a, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

```python

```