### CS102/CS103

Prof. Götz Pfeiffer<br />
School of Mathematics, Statistics and Applied Mathematics<br />
NUI Galway

# Computer Lab 8

## Multiple Choice

For each of the following questions, choose `'a'`, `'b'`, `'c'`, or `'d'` as answer
and type it into the cell below (with quotes).

* What `python` reserved word starts a class definition?

(a) `def`<br />
(b) `class`<br />
(c) `object`<br />
(d) `__init__`<br />

<!-- grade_id: cell-9b270775b2a05bea, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-e8b06319a0c3709e, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```

* A method definition with four formal parameters is generally called with how many actual
parameters?

(a) three<br />
(b) four<br />
(c) five<br />
(d) it depends<br />

<!-- grade_id: cell-a6cf517acdfa6872, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-3f36269f270f2785, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```

* A method definition is similar to a

(a) loop<br />
(b) module<br />
(c) import statement<br />
(d) function definition<br />

<!-- grade_id: cell-5af8c1d097f95db7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-5ae594f28fe0a544, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```

* Within a method definition, the instance variable `x` could be accessed via which expression?

(a) `x`<br />
(b) `self.x`<br />
(c) `self[x]`<br />
(d) `self.getX()`<br />

<!-- grade_id: cell-fd6a8f32cdfcb203, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-079228c64ac1a608, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```

* The term applied to hiding details inside class definitions is

(a) obscuring<br />
(b) subclassing<br />
(c) documentation<br />
(d) encapsulation<br />

<!-- grade_id: cell-3a59aeed108ad967, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-4b5c6877ea3a30d4, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```

```python

```