### CS102/CS103

Prof. Götz Pfeiffer<br />
School of Mathematics, Statistics and Applied Mathematics<br />
NUI Galway

# Computer Lab 8

##  Programming Exercises

In this part of the practical you will write some `python` classes.

In each class definition, the **first line of
the function body** should be **a string that briefly explains the
purpose** of that function.

In some places, a **partial answer** together with the text
```
# YOUR CODE HERE
raise NotImplementedError()
```
is already given in the answer cell.  **Delete** those two lines
and replace them with your own code to complete the answer.

* (7 points) Disks and circles are common geometric objects used in many graphic applications to produce round or circular shapes. For these applications, it is important to design efficient data structures and algorithms to compute with geometric
objects and their properties.  A circular disk in the $x,y$-plane, for example, has a radius, and 
and a centre with $x$- and $y$-coordinates.  Define a class for `Disk` objects, representing
closed circular disks in the plane, that have **instance variables** 
`radius`, `centre_x` and `centre_y`, and the following **methods**:

  * `__init__(self, r, x, y)`, assigning `x` to the instance variable `radius`, and `x` and `y` to
  the centre coordinates; in case the value for `radius` is negative, the method should execute the statement
  ```
  raise ValueError("radius must not be negative")
  ```
  (a radius value of $0$ is allowed).
  
  * `__repr__(self)`, for representing a disk object as a string: a disk with radius $2$ and centre coordinates $(3, 4)$
  should result in the string `"Disk(2, 3, 4)"`, like the expression that was used to create it.
  
  * `diameter(self)`, for computing and returning the diameter of a disk;

  * `circumference(self)`, for computing and returning the diameter of a disk;

  * `area(self)`, for computing and returning the diameter of a disk;
  
  * `distance(self, other)`, for computing the distance between a disk `disk1`
  and another disk `disk2` as `disk1.distance(disk2)`.  Recall that the distance
  between two points $(x_1, y_1)$ and $(x_2, y_2)$ in the plane is given by the
  formula
  $$
  \sqrt{(x_2 - x_1)^2 + (y_2 - y_1)^2}.
  $$
  The distance between two disks can be computes as the distance between their centre points,
  minus the sum of their radii (the plural form of radius).
  
  * Finally, provide an implementation of the method `__mul__(self, other)` so that the expression
  `disk1 * disk2` evaluates to `True` or `False`, depending on whether the disks `disk1` and `disk2`
  overlap or not.



<!-- grade_id: cell-206d1a6af4eecb3d, points: 0, grade: false, solution: true -->
```python
from math import sqrt, pi

class Disk:
    "a disk has a radius and a centre with x,y-coordinates"
    
    def __init__(self, r, x, y):
        "Disk constructor"
        # YOUR CODE HERE
        raise NotImplementedError()

    def __repr__(self):
        return "Disk({}, {}, {})".format(self.radius, self.centre_x, self.centre_y)
    
    def diameter(self):
        # YOUR CODE HERE
        raise NotImplementedError()

    def area(self):
        # YOUR CODE HERE
        raise NotImplementedError()

    def circumference(self):
        # YOUR CODE HERE
        raise NotImplementedError()

    def distance(self, other):
        # YOUR CODE HERE
        raise NotImplementedError()

    def __mul__(self, other):
        "do disks self and other overlap?"
        # YOUR CODE HERE
        raise NotImplementedError()
```

Suppose that `d1`, `d2` and `d3` are disks of radius $2$, one centred at the origin, one at $(1, 0)$ and one at $(3, 4)$.
Then the disks have diameter $4$, and area and circumference $4\pi$. The distance between disks `d1` and `d3` is $1$, disks `d1` and `d2` do overlap, while disks `d1` and `d3` don't. Verify that this is what your methods compute.

```python
d1, d2, d3 = Disk(2, 0, 0), Disk(2, 1, 0), Disk(2, 3, 4)
d1.diameter(), d1.area(), d1.circumference(), d1.distance(d3), d1 * d2, d1 * d3
```

<!-- grade_id: cell-a2fe420d165c96ae, points: 1, grade: true, solution: false -->
```python
"""Check instance variables of Disk objects"""
disk = Disk(1, 2, 3)
assert disk.radius == 1
assert disk.centre_x == 2
assert disk.centre_y == 3

"""Check that negative radius raises ValueError"""
try:
    Disk(-1, 0, 0)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")
```

<!-- grade_id: cell-d9fd4e608b66d6ec, points: 1, grade: true, solution: false -->
```python
"""Check string representation of a disk"""
assert "".join(repr(Disk(1, 2, 3)).split()) == "Disk(1,2,3)"
```

<!-- grade_id: cell-dc9d6adac0c6168d, points: 1, grade: true, solution: false -->
```python
"""Check diameter"""
d1, d2 = Disk(1, 0, 0), Disk(2, 0, 0)
assert d1.diameter() == 2
assert d2.diameter() == 4
```

<!-- grade_id: cell-8a2d6d375beab5f9, points: 1, grade: true, solution: false -->
```python
"""Check circumference"""
d1, d2 = Disk(1, 0, 0), Disk(2, 0, 0)
assert d1.circumference() == 2 * pi
assert d2.circumference() == 4 * pi
```

<!-- grade_id: cell-47a227f15fe629f7, points: 1, grade: true, solution: false -->
```python
"""Check area"""
d1, d2 = Disk(1, 0, 0), Disk(2, 0, 0)
assert d1.area() == pi
assert d2.area() == 4 * pi
```

<!-- grade_id: cell-5bd949f64b46ba75, points: 1, grade: true, solution: false -->
```python
"""Check distance"""
d1, d2 = Disk(1, 0, 0), Disk(2, 3, 4)
assert d1.distance(d2) == 2.0
```

<!-- grade_id: cell-6d24f465466fec9c, points: 1, grade: true, solution: false -->
```python
"""Check overlap"""
d1, d2, d3 = Disk(2, 0, 0), Disk(2, 1, 0), Disk(2, 3, 4)
assert d1 * d2
assert not d1 * d3
```

* (9 points) [Modular arithmetic](https://en.wikipedia.org/wiki/Modular_arithmetic) lies at the heart of modern
encryption methods that are used, for example, for secure communication over the internet. It is built on the concept
of congruence classes: the **congruence class** of $37$ modulo $17$, for example, consists of **all** integers
which have the **same remainder** upon division by $17$ (i.e., $3$). Each congruence class modulo $17$ can be represented by the
**unique** integer $r$ with $0 \leq r < 17$, which is the common remainder upon division by $17$ of all elements in the class.  Congruence classes can be added and multiplied, each class has a negative, some have an inverse.
Define a **class** `Mod` whose instances are congruence classes, represented by **instance variables**
`val` (for the value of the **remainder** $r$, like the number $3$ in the example above) and `mod` (for the **modulus**,
like the number $17$ in the example above), and with the following **methods**:

  * `__init__(self, val, mod)` , assigning `val` to the instance variable `val`, and `mod` to
  the instance variable `mod`; in case the value for `mod` is negative, the method should execute the statement
  ```
  raise ValueError("modulus must be a positive integer")
  ```
  (a modulus value of $0$ is not allowed either).

  * `__repr__(self)`, for representing a `Mod` instance as a string: an object with `val` value $2$ and `mod` value $3$
  should result in the string `"Mod(2, 3)"`, just like the expression that was used to create it.
  
  * `__add__(self, other)`, for computing the sum `self + other` of two `Mod` instances;
  whenever `self` and `other` do **not** have the same `mod` values, the method should execute the statement
  ```
  raise TypeError("operands must have same modulus")
  ```

  * `__mul__(self, other)`, for computing the product `self * other` of two `Mod` instances;
  whenever `self` and `other` do **not** have the same `mod` values, the method should execute the statement
  ```
  raise TypeError("operands must have same modulus")
  ```

  * `__neg__(self)`, for computing the negative `-self` of a `Mod` instance;

  * `inverse(self)`, for computing the modular inverse of `self` if it exists; here the extended gcd
  function might be useful; if the inverse does not exist, the method should execute the statement
  ```
  raise ValueError("{} has no inverse mod {}".format(self.val, self.mod))
  ```

  * `__sub__(self, other)`, for computing the difference `self - other` as the sum `self + -other`;

  * `__truediv__(self, other)`, for computing the quotient `self / other` as the product `self * other.inverse()`;
  
  * `__eq__(self, other)`, for evaluating the comparison `self == other` to `True` only if `self` and `other`
  have the same `mod` and `val` values, and to `False` otherwise
  
The implementation of the extended Euclidean Algorithm as discussed in the lectures is provided below.

<!-- grade_id: cell-70560c39e9adab87, points: 0, grade: false, solution: true -->
```python
def egcd(a, b):
    "find integers  x, y  such that  gcd(a, b) = x*a + y*b"
    if b == 0:
        return (1, 0)
    x, y = egcd(b, a % b)
    x, y = y, x - (a // b) * y
    return (x, y)

class Mod:
    "elements of congruence classes mod m"

    def __init__(self, val, mod):
        "constructor"
        # YOUR CODE HERE
        raise NotImplementedError()
        
    def __repr__(self):
        "string representation"
        # YOUR CODE HERE
        raise NotImplementedError()
        
    def __add__(self, other):
        "self + other"
        # YOUR CODE HERE
        raise NotImplementedError()
        
    def __mul__(self, other):
        "self * other"
        # YOUR CODE HERE
        raise NotImplementedError()
        
    def __neg__(self):
        "-self"
        # YOUR CODE HERE
        raise NotImplementedError()
        
    def inverse(self):
        "Mod(1, m)/self: extended gcd"
        # YOUR CODE HERE
        raise NotImplementedError()
        
    def __sub__(self, other):
        "self - other"
        # YOUR CODE HERE
        raise NotImplementedError()
        
    def __truediv__(self, other):
        "self / other"
        # YOUR CODE HERE
        raise NotImplementedError()
        
    def __eq__(self, other):
        "self == other"
        # YOUR CODE HERE
        raise NotImplementedError()
```

Suppose that `a` is $111$ modulo $1234$ and that `b` is $11111$ modulo $1234$.
Then `a + b` is $116$, `a - b` is $106$, `a * b` is $555$, ` a / b` is $269$ (modulo $1234$),
and `a == b` is `False` (of course).
Verify that this is what your methods compute.

```python
a, b = Mod(111, 1234), Mod(11111, 1234)
a + b, a - b, a * b, a / b, a == b
```

<!-- grade_id: cell-c0ff5a5a78384f4e, points: 1, grade: true, solution: false -->
```python
"""Check instance variables of Mod objects"""
a = Mod(37, 17)
assert a.val == 3
assert a.mod == 17

"""Check that a negative modulus raises ValueError"""
try:
    Mod(37, -17)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")
```

<!-- grade_id: cell-3fd24a7730d3bd72, points: 1, grade: true, solution: false -->
```python
"""Check string representation of a Mod instance"""
assert "".join(repr(Mod(2, 3)).split()) == "Mod(2,3)"
```

<!-- grade_id: cell-08f762f47e6bcc51, points: 1, grade: true, solution: false -->
```python
"""Check sum"""
a, b = Mod(111, 1234), Mod(11111, 1234)
c = a + b
assert c.mod == a.mod
assert c.val == (111 + 11111) % 1234

"""Check that different moduli raise TypeError"""
try:
    Mod(1, 2) + Mod(1, 3)
except TypeError:
    pass
else:
    raise AssertionError("did not raise")
```

<!-- grade_id: cell-5fbdbe21238b39e3, points: 1, grade: true, solution: false -->
```python
"""Check negative"""
a = -Mod(111, 1234)
assert a.mod == 1234
assert a.val == (-111) % 1234
```

<!-- grade_id: cell-5c01c4e3ace11b9b, points: 1, grade: true, solution: false -->
```python
"""Check difference"""
a, b = Mod(111, 1234), Mod(11111, 1234)
c = a - b
assert c.mod == a. mod
assert c.val == (111 - 11111) % 1234
```

<!-- grade_id: cell-6188196585acc297, points: 1, grade: true, solution: false -->
```python
"""Check product"""
a, b = Mod(111, 1234), Mod(11111, 1234)
c = a * b
assert c.mod == a.mod
assert c.val == (111 * 11111) % 1234

"""Check that different moduli raise TypeError"""
try:
    Mod(1, 2) * Mod(1, 3)
except TypeError:
    pass
else:
    raise AssertionError("did not raise")
```

<!-- grade_id: cell-5244632742b94a2d, points: 1, grade: true, solution: false -->
```python
"""Check inverse"""
a = Mod(111, 1234)
c = a.inverse()
assert c.mod == a. mod
assert c.val == 189

"""Check that non-existing inverse raises ValueError"""
try:
    Mod(2, 1234).inverse()
except ValueError:
    pass
else:
    raise AssertionError("did not raise")
```

<!-- grade_id: cell-ffde75014b3a9e9a, points: 1, grade: true, solution: false -->
```python
"""Check quotient"""
a, b = Mod(111, 1234), Mod(11111, 1234)
c = a / b
assert c.mod == a. mod
assert c.val == 269
```

<!-- grade_id: cell-951754bcd7da5e94, points: 1, grade: true, solution: false -->
```python
"""Check equality"""
a, b = Mod(111, 1234), Mod(11111, 1234)
assert a == a
assert a != b
```

```python

```