### CS102/CS103

Prof. Götz Pfeiffer<br />
School of Mathematics, Statistics and Applied Mathematics<br />
NUI Galway

# Computer Lab 1

## True or False

Each of the following statements is either true or false.
Delete the current content of the cell below the statement.
Then type in your answer as either `True` or `False`
(with a captial `T` or `F`).

* A computer is a universal machine.

<!-- grade_id: cell-b9c2056c23075e9d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-3ec807622b11a458, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* Computer Science is the study of computers.

<!-- grade_id: cell-c1b41a7a519fdc91, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-ba83cceb4476ccc9, points: 1, grade: true, solution: false -->
```python

```

* `Python` does not distinguish between integers and floating point numbers.

<!-- grade_id: cell-0776705030eabbce, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-960065372176245b, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* An integer does not have a decimal dot.

<!-- grade_id: cell-f739991c66fa39d0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-0640340d7a0cd2b8, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* In `Python` 3, the quotient of two integers is always an integer.

<!-- grade_id: cell-290fe0e2f1d6746f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-9ff8766d990d5ff9, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* If a string is surrounded by single quotes then the single quotes form part of the string.

<!-- grade_id: cell-a742c69d4a6489aa, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-651611730b6a9ccb, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```