### CS102/CS103

Prof. Götz Pfeiffer<br />
School of Mathematics, Statistics and Applied Mathematics<br />
NUI Galway

# Computer Lab 1

## Multiple Choice

For each of the following questions, choose `'a'`, `'b'`, `'c'`, or `'d'` as answer
and type it into the cell below (with quotes).

What is the fundamental question of computer science?

(a) How fast can a computer compute?<br />
(b) What can be computed?<br />
(c) What is the most efficient programming language?<br />
(d) How much money can a programmer make?<br />


<!-- grade_id: cell-0ed9f15e2a43433b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-bcd9a91832be4ff6, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```

The sum of the integer `11000` and the floating point number `22.33` in `Python`
produces

(a) an error message, indicating that the summands have different types,<br />
(b) an integer like `11022`,<br />
(c) a floting point value like `11022.33`,<br />
(d) nothing - the `Python` interpreter freezes.<br />

<!-- grade_id: cell-fb32b0afb00ac905, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-8d20de70da81a5af, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```