### CS102/CS103

Prof. Götz Pfeiffer<br />
School of Mathematics, Statistics and Applied Mathematics<br />
NUI Galway

# Computer Lab 1

##  Programming Exercises

Follow the instructions. Replace the content of each answer cell with your own answer.  At the end, each answer cell should contain **exactly one line** of `Python` code.  You can use the cells below each answer box to check if your answer has the correct format.  **Except for the last one**, which should be **two lines**: a comment and a line of `Python`.

* Use the cell below to experiment and observe the effects of the following commands one at a time.

```python
print("Hello, world!")
print("Hello", "world!")
"Hello " + "world!"
print 3
print 3.00
2 + 3
2.0 + 3
2.0 + 3.0
"2" + "3"
2 * 3
2 ** 3
3 / 2
3 // 2
3 % 2
3.0 / 2.0
3.0 // 2.0
```

<!-- grade_id: cell-8f851c7cb4c045c9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3dcc80bf8ba7455f, points: 1, grade: true, solution: false -->
```python
_input = list(filter(lambda x: x[:3] != '###', _i.split('\n')))
assert(len(_input) == 1)
```

* Write a `Python` program that computes the product of $1234$ and $5678$.

<!-- grade_id: cell-754f904074ac2077, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ad76dba0b108be27, points: 1, grade: true, solution: false -->
```python
_input = list(filter(lambda x: x[:3] != '###', _i.split('\n')))
_answer = _
assert(len(_input) == 1)
assert(type(_answer) == int)
```

* Write a `Python` program that computes the quotient of the sum of
$445.5$ and $1600.4$ and the product of $34.1$ and $6.2$.

<!-- grade_id: cell-120fa78d0daa4223, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c69145c43ccbdf22, points: 1, grade: true, solution: false -->
```python
_input = list(filter(lambda x: x[:3] != '###', _i.split('\n')))
_answer = _
assert(len(_input) == 1)
assert(type(_answer) == float)
```

* Write a `Python` program that determines the number of **full** weeks left until the end of the year, if there are $94$ days remaining.

<!-- grade_id: cell-571b54d9be8ef581, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f82ec7d8e2c7a2f5, points: 1, grade: true, solution: false -->
```python
_input = list(filter(lambda x: x[:3] != '###', _i.split('\n')))
_answer = _
assert(len(_input) == 1)
assert(type(_answer) == int)
```

* Write a `Python` that computes the sum of $10$, $14$ and $18$ and reports back by writing the text `'The answer is'` (without the quotes) followed by the sum, followed by a full stop on the screen.

<!-- grade_id: cell-94422051442aacfa, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b856bd9a029b4843, points: 1, grade: true, solution: false -->
```python
_input = list(filter(lambda x: x[:3] != '###', _i.split('\n')))
assert(len(_input) == 1)
```

* Write a `Python` program that allows you to determine whether $200001079$ is a multiple of $5323$.

<!-- grade_id: cell-f86392c630f3d5ba, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-82c92b610e4f004c, points: 1, grade: true, solution: false -->
```python
_input = list(filter(lambda x: x[:3] != '###', _i.split('\n')))
_answer = _
assert(len(_input) == 1)
assert(type(_answer) == int)
```

* Write a `Python` program that prints $15$ letters `'y'` (without the quotes) followed by $15$ letters `'e'` followed by $15$ letters `'s'` followed by $15$ `'!'`.

<!-- grade_id: cell-0f580e74d04881aa, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6f8185d8c04eed5e, points: 1, grade: true, solution: false -->
```python
_input = list(filter(lambda x: x[:3] != '###', _i.split('\n')))
assert(len(_input) == 1)
```

* Write a `Python` program that determines the number of seconds in a day.

<!-- grade_id: cell-e7dc887db8f200fa, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-cefe711a9b5c335d, points: 1, grade: true, solution: false -->
```python
_input = list(filter(lambda x: x[:3] != '###', _i.split('\n')))
_answer = _
assert(len(_input) == 1)
assert(type(_answer) == int)
```

* Write a `Python` program that determines the number of seconds in a week.

<!-- grade_id: cell-25e32e96b0a03066, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-292daa6832783cb6, points: 1, grade: true, solution: false -->
```python
_input = list(filter(lambda x: x[:3] != '###', _i.split('\n')))
_answer = _
assert(len(_input) == 1)
assert(type(_answer) == int)
```

* Write a `Python` program that computes your age in seconds.

<!-- grade_id: cell-c98898e5af2acfa3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-eeb27a13ba768062, points: 1, grade: true, solution: false -->
```python
_input = list(filter(lambda x: x[:3] != '###', _i.split('\n')))
_answer = _
assert(len(_input) == 1)
assert(type(_answer) == int)
```

* Write a `Python` program that computes the interest on an investment of EUR 89235.41 at an interest rate of 4.5 percent.

<!-- grade_id: cell-91b7108cbb34b519, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8fff4e65b1ada890, points: 1, grade: true, solution: false -->
```python
_input = list(filter(lambda x: x[:3] != '###', _i.split('\n')))
_answer = _
assert(len(_input) == 1)
assert(type(_answer) == float)
```

* Write a `Python` program that determines the future value of a sum of EUR 189235.41 invested for two years at an interest rate of 4.4 percent per year, with interest compunded every three months.

<!-- grade_id: cell-513ce0b06f66f464, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-366ae86494b39cd4, points: 1, grade: true, solution: false -->
```python
_input = list(filter(lambda x: x[:3] != '###', _i.split('\n')))
_answer = _
assert(len(_input) == 1)
assert(type(_answer) == float)
```

* Using a value of $3.14159$ for the constant $\pi$, write a `Python` program that computes the area of a circle of radius 12.3 meter (in square meters).

<!-- grade_id: cell-f6f6f0881a8aefa0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b766eb0c64221203, points: 1, grade: true, solution: false -->
```python
_input = list(filter(lambda x: x[:3] != '###', _i.split('\n')))
_answer = _
assert(len(_input) == 1)
assert(type(_answer) == float)
```

* Think of a task that you now trust `Python` can perform, describe it in a one-line comment, and write a (one-line) `Python` program that implements the task.

<!-- grade_id: cell-b1b8f6bb362bae4e, points: 0, grade: false, solution: true -->
```python
#  <-- Describe task in one line here ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-35f5f39c669b6f7d, points: 1, grade: true, solution: false -->
```python
_input = list(filter(lambda x: x[:3] != '###', _i.split('\n')))
assert(len(_input) == 2)
```