### CS102/CS103

Prof. Götz Pfeiffer<br />
School of Mathematics, Statistics and Applied Mathematics<br />
NUI Galway

# Computer Lab 4

## True or False

Each of the following statements is either true or false.
Delete the current content of the cell below the statement.
Then type in your answer as either `True` or `False`
(with a capital `T` or `F`).

* A simple decision can be implemented with an `if` statement.

<!-- grade_id: cell-b0ac1b5fe2bce65e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-d6f514f005a88682, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* In `python` conditions, $\neq$ is written as `/=`.

<!-- grade_id: cell-c4e154b52dcdc7e6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-18bbba2dd15224a4, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* A two-way decision is implemented using an `if`-`elif` statement.

<!-- grade_id: cell-5035797c7e4cf684, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-5fa283706b04730d, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* The `math.sqrt()` function cannot compute the square root of a negative number.

<!-- grade_id: cell-fac64e4052d362e7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-d84e6daca4ef3a30, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* The condition `x <= y <= z` is allowed in `python`.

<!-- grade_id: cell-8af7dfcaa9a70caa, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-3acf5ed91996a628, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* The `while` statement in `python` implements a definite loop.

<!-- grade_id: cell-e5c7e2165a400691, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-222d116c0c45bb99, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* A counted loop is a special case of a definite loop.

<!-- grade_id: cell-181c3c8d2bdb9ab1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-bf906a283c322735, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

```python

```