### CS102/CS103

Prof. Götz Pfeiffer<br />
School of Mathematics, Statistics and Applied Mathematics<br />
NUI Galway

# Computer Lab 4

## Multiple Choice

For each of the following questions, choose `'a'`, `'b'`, `'c'`, or `'d'` as answer
and type it into the cell below (with quotes).

* A statement that controls the execution of other statements is called 

(a) a boss structure<br />
(b) a super structure<br />
(c) a control structure<br />
(d) a branch<br />

<!-- grade_id: cell-1c2493ecb68b6f6a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-a85f9f679f608e18, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```

* An expression that evaluates to either true or false is called

(a) operational<br />
(b) boolean<br />
(c) simple<br />
(d) compound<br />

<!-- grade_id: cell-f8f7aca93325b6ba, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-3823d12c646af2c9, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```

* The literals for type `bool` in `python` are

(a) `T`, `F`<br />
(b) `True`, `False`<br />
(c) `true`, `false`<br />
(d) `0`, `1`<br />

<!-- grade_id: cell-6600d187f575a0a7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-ba81c8839785e64a, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```

* In `python` the body of a decision statement is indicated by

(a) indentation<br />
(b) parentheses<br />
(c) curly braces<br />
(d) a colon<br />

<!-- grade_id: cell-38e13b7d0801e81e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-c53189a79941a3f8, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```

```python

```