### CS102/CS103

Prof. Götz Pfeiffer<br />
School of Mathematics, Statistics and Applied Mathematics<br />
NUI Galway

# Computer Lab 4

##  Programming Exercises

In this part of the practical you will write some simple `python` functions.

In each function definition, the **first line of
the function body** should be **a string that briefly explains the
purpose** of that function.

Some of these are **function versions** of the **minimal IPO** programs from
Lab 2.  In each such case it takes again **exactly 4 lines of code** to 
write such a function:

1. (I) the **function definition heading** which declares the **input** variable(s)
as parameter(s).

2. A **string** that briefly explains the
purpose of this function.

3. (P) A **statement** that assigns the value of a expression involving the input
variable to an output variable.

4. (O) A **return statement** involving the output variable.

Here is the function version of the temperature converter in  this format:

```python
def f2c(fahrenheit):
    "this function converts Fahrenheit into Celsius"
    celsius = (fahrenheit - 32) * 5 / 9
    return celsius
```

We already know that 77 degress Fahrenheit correspond to 25 degrees Celsius.  The function
`f2c()` hopefully confirms this:

```python
f2c(77)
```

In some of the exercises, a **partial answer** together with the text
```
# YOUR CODE HERE
raise NotImplementedError()
```
is already given in the answer cell.  **Delete** those two lines
and replace them with your own code to complete the answer.

* Write a `python` function `c2f()` that converts Celsius into Fahrenheit

<!-- grade_id: cell-2faf0e81992f1428, points: 0, grade: false, solution: true -->
```python
def c2f(celsius):
    "Convert Celsius into Fahrenheit"
    # YOUR CODE HERE
    raise NotImplementedError()
```

25 degrees Celsius correspond to 77 degrees Fahrenheit.  Verify that this is what your function
computes.

```python
c2f(25)
```

<!-- grade_id: cell-bdadbd52df619895, points: 2, grade: true, solution: false -->
```python
"""Check that c2f returns the correct answer for various inputs"""
assert c2f(20.0) == 68.0
assert c2f(25.0) == 77.0
assert c2f(30.0) == 86.0
assert c2f(35.0) == 95.0

"""Check that the function has a docstring"""
assert c2f.__doc__ != None
```

* Write a `python` function `miles2km()` that converts distances measured in miles into kilometres.
(Hint: [wikipedia](http://en.wikipedia.org/wiki/Mile))

<!-- grade_id: cell-c37f938615d68179, points: 0, grade: false, solution: true -->
```python
def miles2km(miles):
    # YOUR CODE HERE
    raise NotImplementedError()
```

10 miles are 16.09344 miles. Verify that this is what your function computes.

```python
miles2km(10)
```

<!-- grade_id: cell-d5d709fd1e9a01c8, points: 2, grade: true, solution: false -->
```python
"""Check that miles2km returns the correct answer for various inputs"""
assert miles2km(0.1) - 0.1609344 < 0.01
assert miles2km(1) - 1.609344 < 0.1
assert miles2km(10) - 16.09344 < 1

"""Check that the function has a docstring"""
assert miles2km.__doc__ != None
```

* Write a `python` function `mpg2lp100km()` that converts miles per gallon into litres per 100 kilometres.

<!-- grade_id: cell-a56b982564774536, points: 0, grade: false, solution: true -->
```python
def mpg2lp100km(mpg):
    # YOUR CODE HERE
    raise NotImplementedError()
```

35 miles per gallon are roughly 8.1 litres per 100 km. Verify that this is what your function computes.

```python
mpg2lp100km(40)
```

<!-- grade_id: cell-dbcc2bc08814bd56, points: 2, grade: true, solution: false -->
```python
"""Check that mpg2lp100km returns the correct answer for various inputs"""
assert mpg2lp100km(40) - 7.1 < 0.1
assert mpg2lp100km(35) - 8.1 < 0.1
assert mpg2lp100km(30) - 9.4 < 0.1

"""Check that the function has a docstring"""
assert mpg2lp100km.__doc__ != None
```

* Write a `python` function `imperial2metric()` that converts a length measured in feet and inches
into centimetres. (Hint. [wikipedia](http://en.wikipedia.org/wiki/Imperial_units))

<!-- grade_id: cell-53ff98b39f8cecd8, points: 0, grade: false, solution: true -->
```python
def imperial2metric(feet, inches):
    # YOUR CODE HERE
    raise NotImplementedError()
```

8 feet and 4 inches are 254cm. Verify that this is what your function computes.

```python
imperial2metric(8, 4)
```

<!-- grade_id: cell-7d44f366728e3a5c, points: 2, grade: true, solution: false -->
```python
"""Check that imperial2metric returns the correct answer for various inputs"""
assert imperial2metric(1, 0) == 12 * 2.54
assert imperial2metric(0, 1) == 2.54
assert imperial2metric(8, 4) == 254.0

"""Check that the function has a docstring"""
assert imperial2metric.__doc__ != None
```

The remaining functions to write might require the use of decision statements, and could
therefore be longer than just 4 lines of code.

* A 5-point quiz is graded on the scale 5:A, 4:B, 3:C, 2:D, 0 or 1:F.  Write a `python` function
`grade5()` that converts a quiz score into the corresponding grade.

<!-- grade_id: cell-d6b057dfff8fb869, points: 0, grade: false, solution: true -->
```python
def grade5(score):
    # YOUR CODE HERE
    raise NotImplementedError()
```

So a score of 3 should give a grade `C`. Verify that this is what your function computes.

```python
grade5(3)
```

<!-- grade_id: cell-a739020f655c7dab, points: 2, grade: true, solution: false -->
```python
"""Check that grade5 returns the correct answer for various inputs"""
assert grade5(0) == 'F'
assert grade5(1) == 'F'
assert grade5(2) == 'D'
assert grade5(3) == 'C'
assert grade5(4) == 'B'
assert grade5(5) == 'A'

"""Check that the function has a docstring"""
assert grade5.__doc__ != None
```

A 100-point exam is graded on the scale 90-100:A, 80-89:B, 70-79:C, 60-69:D, 50-59:E, less than 50:F.  Write a `python` function `grade100()` that converts an exam score into the corresponding grade.

<!-- grade_id: cell-17f0e1b0f6748a54, points: 0, grade: false, solution: true -->
```python
def grade100(score):
    # YOUR CODE HERE
    raise NotImplementedError()
```

So a score of 65 should yield a grade of `D`. Verify that this is what your function computes.

```python
grade100(65)
```

<!-- grade_id: cell-a2f431f9a5dc0e1d, points: 2, grade: true, solution: false -->
```python
"""Check that grade100 returns the correct answer for various inputs"""
assert grade100(0) == 'F'
assert grade100(10) == 'F'
assert grade100(49) == 'F'
assert grade100(50) == 'E'
assert grade100(51) == 'E'
assert grade100(59) == 'E'
assert grade100(60) == 'D'
assert grade100(70) == 'C'
assert grade100(80) == 'B'
assert grade100(90) == 'A'
assert grade100(100) == 'A'

"""Check that the function has a docstring"""
assert grade100.__doc__ != None
```

```python

```