### CS102/CS103

Prof. Götz Pfeiffer<br />
School of Mathematics, Statistics and Applied Mathematics<br />
NUI Galway

# Computer Lab 5

## True or False

Each of the following statements is either true or false.
Delete the current content of the cell below the statement.
Then type in your answer as either `True` or `False`
(with a capital `T` or `F`).

* Unlike lists, `python` strings are not mutable.

<!-- grade_id: cell-9746fdbc89a836b8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-15ef3b81b46b456f, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* Items can be removed from a list with the `del` operator.

<!-- grade_id: cell-c8459a3847e6ebab, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-f55143eaf69ab430, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* A `python` typle is similar to a list, except that tuples are immutable.

<!-- grade_id: cell-6c4d0128df36e413, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-3ca5a006da1f4154, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* A `python` dictionary is a kind of unordered collection.

<!-- grade_id: cell-e14ca73f593ec176, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-8c5581749a7a58c5, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* The boolean operator `or` returns `True` only when both its operands are true.

<!-- grade_id: cell-6e7c9033052e1315, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-a292b97a8fceaae9, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* `not (P and Q) == (not P) and (not Q)`

<!-- grade_id: cell-6dad1de60341accf, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-7c27199da7a91754, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* `True or False`

<!-- grade_id: cell-6422d11eea1a79e1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-f610046d7ab7e60a, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

```python

```