### CS102/CS103

Prof. Götz Pfeiffer<br />
School of Mathematics, Statistics and Applied Mathematics<br />
NUI Galway

# Computer Lab 5

## Multiple Choice

For each of the following questions, choose `'a'`, `'b'`, `'c'`, or `'d'` as answer
and type it into the cell below (with quotes).

* Which of the following correctly tests whether `x` is even?

(a) `x // 2 == 0`<br />
(b) `x / 2 == 0`<br />
(c) `x % 2 == 0`<br />
(d) `x % 2 == x`<br />

<!-- grade_id: cell-b9c9c9dd615933e3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-3eb19aff93f05c15, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```

* What keyword parameter is used to send a key-function to the `sort()` method?

(a) `reversed`<br />
(b) `reverse`<br />
(c) `key`<br />
(d) `cmp`<br />

<!-- grade_id: cell-aa9cdb4bc7e94911, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-d39ca1e9841bfdd5, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```

* Which of the following is not a dictionary method?

(a) `get()`<br />
(b) `keys()`<br />
(c) `sort()`<br />
(d) `clear()`<br />

<!-- grade_id: cell-e60c53b8dfb0aa33, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-8b951f1bb5d955ca, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```

* Which line would not be found in a truth table for `and`?

(a) `True  True  True`<br />
(b) `True  False True`<br />
(c) `False True  False`<br />
(d) `False False False`<br />

<!-- grade_id: cell-6dc0c7312ea0df97, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-34f5212cccd3b9e4, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```

```python

```