### CS102/CS103

Prof. Götz Pfeiffer<br />
School of Mathematics, Statistics and Applied Mathematics<br />
NUI Galway

# Computer Lab 5

##  Programming Exercises

In this part of the practical you will write some `python` functions.

In each function definition, the **first line of
the function body** should be **a string that briefly explains the
purpose** of that function.

In some of these functions, a **decision statement** is needed to 
distinguish several cases.  In other functions, a variant of the **accumulator pattern** might be useful.
Here, an **accumulator variable** is first **initialized** with a suitable value,
then during a loop over some sequence, the accumulator variable is **updated in each iteration**
of the loop, and finally the final value of the accumulator variable is returned.
The **loop** in the accumulator can be several loops **nested** inside each other.  

The following `python` progam uses the accumulator pattern to compute the product of a list of numbers
in the accumulator variable `total`.

```python
def product_numbers(numbers):
    "compute the product of a list of numbers"
    total = 1
    for number in numbers:
        total *= number
    return total
```

We know that the product of the numbers $1,2,3$ and $4$ is $24$.  That's what the function should compute.

```python
product_numbers(range(1,5))
```

In some of the exercises, a **partial answer** together with the text
```
# YOUR CODE HERE
raise NotImplementedError()
```
is already given in the answer cell.  **Delete** those two lines
and replace them with your own code to complete the answer.

* A company pays their employees EUR 20 per hour up to forty hours
a week, and EUR 30 for every hour in excess of that.
Write a `python` function `pay()` that takes the number of 
hours worked in a week and computes and returns the pay in euro.
As this function deals with integer amounts of hours
and euro, it should return a value of type `int`.

<!-- grade_id: cell-e27b030570cc817c, points: 0, grade: false, solution: true -->
```python
def pay(hours):
    "Compute weekly pay"
    # YOUR CODE HERE
    raise NotImplementedError()
```

40 hours of work at a rate of EUR 20 per hour thus would earn a pay
of EUR 800.  Verify that this is what your function
computes.

```python
pay(40)
```

<!-- grade_id: cell-bdadbd52df619895, points: 2, grade: true, solution: false -->
```python
"""Check that pay returns the correct answer for various inputs"""
assert pay(40) == 800
assert pay(0) == 0
assert pay(30) == 600
assert pay(50) == 1100
assert pay(39) == 780
assert pay(41) == 830

"""Check that the function has a docstring"""
assert pay.__doc__ != None
```

* In a certain country, a single employee  pays 20% tax 
on an annual income of up to EUR 33,800,
and 40% on every euro earned in excess of that.
Write a Python function `tax()`
that takes the annual income in euro as argument, and
computes and returns the amount (in euro) of tax to be paid on that.

<!-- grade_id: cell-3d4ff229eb796a89, points: 0, grade: false, solution: true -->
```python
def tax(income):
    # YOUR CODE HERE
    raise NotImplementedError()
```

On an annual income of EUR 33,800, one would have to pay EUR 6760 in taxes.
Verify that this is what your function computes.

```python
tax(33800)
```

<!-- grade_id: cell-30b5544156c12f87, points: 2, grade: true, solution: false -->
```python
"""Check that tax returns the correct answer for various inputs"""
assert tax(33800) == 33800 * 0.20
assert tax(43800) == 33800 * 0.20 + 10000 * 0.40
assert tax(23800) == 23800 * 0.20
assert tax(0) == 0 * 0.20

"""Check that the function has a docstring"""
assert tax.__doc__ != None
```

*  Write a Python function `three_digits()`
that takes as argument a digit `d` between $0$ and $9$,
and then uses nested for loops to generate a list of all three digit numbers
such that the first digit is at least `d`,
the second digit is as least as large as the first,
and the third digit is at least as large as the second.
Here, it helps to use an **accumulator variable** numbers, initialized
as an empty list.

<!-- grade_id: cell-897ac83b32c498fc, points: 0, grade: false, solution: true -->
```python
def three_digits(d):
    "list all 3-digit numbers under certain conditions."
    # YOUR CODE HERE
    raise NotImplementedError()
```

The numbers satisfying the conditions for $d = 8$ are
$888$, $889$, $899$ and $999$. Verify that this is what your function computes.

```python
three_digits(8)
```

<!-- grade_id: cell-a772ee93458cdb02, points: 2, grade: true, solution: false -->
```python
"""Check that three_digits returns the correct answer for various inputs"""
assert three_digits(9) == [999]
assert three_digits(8) == [888, 889, 899, 999]
assert(len(three_digits(0)) == 220)

"""Check that the function has a docstring"""
assert three_digits.__doc__ != None
```

* Write a `python` function `letter_frequencies()` that takes
a string `text` as argument and computes and returns
a **dictionary** that has as keys all the letters
occuring in the string `text` (and no other keys), which each have as their
value a count of the number of appearances of that letter in the text.  
The function should ignore digits, spaces, punctuation and everything
else that is not a letter.
The function should treat each uppercase letter as the corresponding lower case letter.  The dictionary thus should only have lower case letters as its keys.  
You can use the string method `isalpha()` to decide whether
a symbol from the string is a letter (`'a'.isalpha()` returns `True`
and `'1'.isalpha()` returns `False`).
You can use the string method `lower()` to turn any letter into
its lowercase form (`'A'.lower()` returns `'a'` and `'a'.lower()`
returns `'a'` as well).
Here it helps to use an **accumulator variable** `letters`, initialized as an empty dictionary (`{}`).

<!-- grade_id: cell-ecbc7725e61ac272, points: 0, grade: false, solution: true -->
```python
def letter_frequencies(text):
    # YOUR CODE HERE
    raise NotImplementedError()
```

The string `"Mississippi!"` contains 4 letters `'i'`, 4 letters `'s'`
and 1 of `'m'` and `'p'`. Verify that this is what your function computes.

```python
letter_frequencies("Mississipi!")
```

<!-- grade_id: cell-dbcc2bc08814bd56, points: 2, grade: true, solution: false -->
```python
"""Check that letter_frequencies returns the correct answer for various inputs"""
assert letter_frequencies("Mississipi!") == {'i': 4, 'm': 1, 'p': 1, 's': 4}
assert letter_frequencies("A B C\n") == {'a': 1, 'b': 1, 'c': 1}
assert letter_frequencies("(1) A,B-C.") == {'a': 1, 'b': 1, 'c': 1}
assert letter_frequencies("") == {}

"""Check that the function has a docstring"""
assert letter_frequencies.__doc__ != None
```

* Now write a `python` function `add_letter_frequencies(letters, text)` that takes a (possibly empty)
dictionary `letters` of letter counts and a string `text` as arguments counts the letters (as lowercase letters)
in `text` as above, and **adds** those counts to the ones already in the dictionary `letters`.

<!-- grade_id: cell-cc3eef2acc2ca14b, points: 0, grade: false, solution: true -->
```python
def add_letter_frequencies(letters, text):
    "add letter frequencies in text to letters dict"
    # YOUR CODE HERE
    raise NotImplementedError()
```

The dictionary `{'a': 2, 'b': 1}` and the text `"abc"` should result in the dictionary `{'a': 3, 'b': 2, 'c': 1}`.
Verify that this is what your function computes.

```python
add_letter_frequencies({'a': 2, 'b': 1}, "abc")
```

<!-- grade_id: cell-6a71a30d89b7e77f, points: 2, grade: true, solution: false -->
```python
"""Check that add_letter_frequencies returns the correct answer for various inputs"""
assert add_letter_frequencies({}, "Mississipi!") == {'i': 4, 'm': 1, 'p': 1, 's': 4}
assert add_letter_frequencies({}, "(1) A,B-C.\n") == {'a': 1, 'b': 1, 'c': 1}
assert add_letter_frequencies({}, "") == {}
assert add_letter_frequencies({'a': 2, 'b': 1}, "abc") == {'a': 3, 'b': 2, 'c': 1}

"""Check that the function has a docstring"""
assert add_letter_frequencies.__doc__ != None
```

* Write a `python` function `letter_frequencies_story(filename)` that takes a file name as argument,
opens the file, then reads it line by line, and uses the 
`add_letter_frequencies()` function from above to update an accumulator dictionary with the 
letter counts of each line.
The function should return a dictionary of letter frequencies in the given file.

<!-- grade_id: cell-63adcd2dd0c5cb12, points: 0, grade: false, solution: true -->
```python
def letter_frequencies_story(filename):
    "count letters in a file"
    # YOUR CODE HERE
    raise NotImplementedError()
```

The file `scan.txt` contains the short story _A Scandal in Bohemia_ by Arthur Conan Doyle.  There are exactly $46$ letters `'x'` in this story. Verify that this is what your function computes.

```python
letter_frequencies_story("scan.txt")['e']
```

<!-- grade_id: cell-b9b224444e50dd61, points: 2, grade: true, solution: false -->
```python
"""Check that letter_frequencies_story returns the correct answer for various inputs"""
assert letter_frequencies_story("scan.txt")['e'] == 4443

"""Check that the function has a docstring"""
assert letter_frequencies_story.__doc__ != None
```

* Write a function `second(values)` that returns the **second** item in a list `values`, i.e., the item at
position $1$.

<!-- grade_id: cell-9f8377f19c836b4f, points: 0, grade: false, solution: true -->
```python
def second(values):
    # YOUR CODE HERE
    raise NotImplementedError()
```

The second item in the sequence `"abc"` is the letter `'b'`. Verify that this is what your function computes.

```python
second("abc")
```

<!-- grade_id: cell-a5c717db57a44f7d, points: 2, grade: true, solution: false -->
```python
"""Check that second returns the correct answer for various inputs"""
assert second("abc") == 'b'
assert second(range(4)) == 1

```

* **Bonus Question**:  Finally, write a `python` function `print_frequencies(counts, n)` that prints a 
nice report on the most frequent `n` letters in a dictionary `counts` of letter counts.
This function should: 
    1. Use the dictionary method `items()` and the `list()` function to produce a
**list** `items` of all the key-value pairs in the dictionary.
    2. Use the `sort()` method to sort the list items
with **keyword parameters** `key` set to the above function
`second`, and `reverse` set to `True`.
    3. Loop over the first `n` items in the list, and for each item,
print the letter and the corresponding count.  One can use
string formatting on the format string `"{0:<2}{1:>7}"` to 
obtain a nicely aligned listing.

(This bonus question is self-rewarding: no points can be earned, or lost.)

<!-- grade_id: cell-ca5ea0aeeada14e6, points: 0, grade: false, solution: true -->
```python
def print_frequencies(counts, n):
    "print a list of frequencies"
    # YOUR CODE HERE
    raise NotImplementedError()
```

The report on the letter frequencies in the file `scan.txt` looks like:
```
e    4443
t    3213
a    2884
```
What does your program print as the top ten?

```python
letters = letter_frequencies_story("scan.txt")
print_frequencies(letters, 10)
```

<!-- grade_id: cell-7388e6968cb8d250, points: 0, grade: true, solution: false -->
```python
"""Check that the function has a docstring"""
assert print_frequencies.__doc__ != None
```

```python

```