### CS102/CS103

Prof. Götz Pfeiffer<br />
School of Mathematics, Statistics and Applied Mathematics<br />
NUI Galway

# Computer Lab 2

## True or False

Each of the following statements is either true or false.
Delete the current content of the cell below the statement.
Then type in your answer as either `True` or `False`
(with a captial `T` or `F`).

* Names (identifiers) in a `python` program must start with a letter or an underscore (`_`).

<!-- grade_id: cell-73f48fe9d723d2b1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-cbd0538cd1659bda, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* In `python`, `x = x + 1` is a legal statement.

<!-- grade_id: cell-fac1a561d0c77951, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-a6d235127c6f1743, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* A variable is used to give a name to a value so it can be referred to in other places.

<!-- grade_id: cell-7249bb6b69ddb7f4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-1c5cccaaacf8695d, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* A loop is used to skip over a section of a program.

<!-- grade_id: cell-7b38d02a2e150dde, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-0fcd6e1c3bb16afc, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```