### CS102/CS103

Prof. Götz Pfeiffer<br />
School of Mathematics, Statistics and Applied Mathematics<br />
NUI Galway

# Computer Lab 2

## Multiple Choice

For each of the following questions, choose `'a'`, `'b'`, `'c'`, or `'d'` as answer
and type it into the cell below (with quotes).

* What is the correct formula for converting Celsius ($C$) into Fahrenheit ($F$)?

(a) $F = \tfrac95C + 32$<br />
(b) $F = \tfrac59C - 32$<br />
(c) $F = B^2 - 4AC$<br />
(d) $F = \tfrac{212 - 32}{100 - 0}$<br />

<!-- grade_id: cell-7f3b0a8aadae3055, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-e6dfd6ed30f5cb84, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```

* Which of the following is not allowed as a name of a variable in `python`?

(a) `text`<br />
(b) `tExt`<br />
(c) `2text`<br />
(d) `text4U`<br />

<!-- grade_id: cell-71b89bc951301d9d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-79e46bbdf1a93e4e, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```

* Which of the following are not used in `python` expressions?

(a) variables<br />
(b) statements<br />
(c) operators<br />
(d) integers<br />

<!-- grade_id: cell-71f482cfef68d681, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-8156977a30c7d5bb, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```

* A fragment of code that produces or calculates a new data value is called

(a) variable<br />
(b) expression<br />
(c) productive clause<br />
(d) assignment statement<br />

<!-- grade_id: cell-2b84fca4c67e4ddf, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-561788102c2888e4, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```

* For the following copy of the chaos program, select **all the statements** below which are true,
and submit their letters in one string as your answer.  (That is `"abc"` if you think statements (a), (b) 
and (c) are true, and all other statements are false.)
```python
# A simple program illustrating chaotic behavior
x = 0.25
for i in range(10):
    x = 3.9 * x * (1 - x)
    print(x)
```

(a) The only variable in this program is `x`.<br />
(b) The line `for i in range(10):` is a complete statement.<br />
(c) The line starting with `#` is a comment.<br />
(d) There are exactly two assigment statements in this program.<br />
(e) The code fragment `range(10)` is an expression.<br />
(f) The program would work in the same way if the last two lines were  not indented by spaces.<br />

<!-- grade_id: cell-4b82b99f9303984a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-70b9ecd58f8642e8, points: 3, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```