### CS102/CS103

Prof. Götz Pfeiffer<br />
School of Mathematics, Statistics and Applied Mathematics<br />
NUI Galway

# Computer Lab 2

##  Programming Exercises

In this part of the practical you will write some simple `python` programs in **minimal IPO format**,
designed to work smoothly with the `python` cells in the notebook.
Such a program consists of **exactly 4 lines**:

1. A **comment**, describing what the program does.
2. (I) An **input value** is assigned to a variable;
3. (P) The value is **processed** and the result is assigned to another variable;
4. (O) The variable containing the result on its own, to print that value below the cell.

You can use the cells below each answer box to check if your answer has the correct format.  

Here is a copy of the temperature converter in minimal IPO format:

```python
# This cell contains code to convert Fahrenheit into Celsius
fahrenheit = 77
celsius = (fahrenheit - 32) * 5 / 9
celsius
```

Executing the cell will convert the given Fahrenheit temperature int Celsius and print
the resulting value below the cell.  To convert a different Fahrenheit temperature, simply
replace the given value in the second line.  What is 69F in Celsius? (No points for the correct answer.)

* In the cell below, write a `python` program that converts Celsius into Fahrenheit.

<!-- grade_id: cell-a5eb5dfa0d18d39f, points: 0, grade: false, solution: true -->
```python
# This cell contains code to convert Celsius into Fahrenheit
celsius = 25
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to check if your code has the required 4 lines.)

<!-- grade_id: cell-a418aa9c8b8aa460, points: 2, grade: true, solution: false -->
```python
_input = list(filter(lambda x: x[:3] != '###', _i.split('\n')))
assert(len(_input) == 4)
```

* In the cell below, write a `python` program that converts distances measured in kilometres to miles.
(Hint: [wikipedia](http://en.wikipedia.org/wiki/Mile))

<!-- grade_id: cell-ac196efba8bd8713, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to check if your code has the required 4 lines.)

<!-- grade_id: cell-bf618229ee7b7046, points: 2, grade: true, solution: false -->
```python
_input = list(filter(lambda x: x[:3] != '###', _i.split('\n')))
assert(len(_input) == 4)
```

Here is a copy of the average calculator in minimal IPO format

```python
# This cell contains code to compute the average of two exam scores
score1, score2 = 96, 99
average = (score1 + score2) / 2
average
```

* In the cell below, write a `python` program that computes the average of **three** exam scores.

<!-- grade_id: cell-b9029bbf5d1badea, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to check if your code has the required 4 lines.)

<!-- grade_id: cell-09ee198283484667, points: 2, grade: true, solution: false -->
```python
_input = list(filter(lambda x: x[:3] != '###', _i.split('\n')))
assert(len(_input) == 4)
```

* In the cell below, write a `python` program that converts a length measured in feet and inches into centimetres.
Use simultaneous assignment for the values of `feet` and `inches`.
(Hint. [wikipedia](http://en.wikipedia.org/wiki/Imperial_units))

<!-- grade_id: cell-909ad4a4a7e98bc2, points: 0, grade: false, solution: true -->
```python
# This cell contains a program that converts feet and inches into centimetres.
feet, inches = 6, 4  # six foot 4 ...
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to check if your code has the required 4 lines.)

<!-- grade_id: cell-cd93c0e1b64ceda2, points: 2, grade: true, solution: false -->
```python
_input = list(filter(lambda x: x[:3] != '###', _i.split('\n')))
assert(len(_input) == 4)
```

* In the cell below, write a `python` program that converts miles per gallon into litres per 100 kilometres.

<!-- grade_id: cell-ce7f5905a870be65, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to check if your code has the required 4 lines.)

<!-- grade_id: cell-8a93881a9d0acb4f, points: 2, grade: true, solution: false -->
```python
_input = list(filter(lambda x: x[:3] != '###', _i.split('\n')))
assert(len(_input) == 4)
```

```python

```