### CS102/CS103

Prof. Götz Pfeiffer<br />
School of Mathematics, Statistics and Applied Mathematics<br />
NUI Galway

# Computer Lab 3

## True or False

Each of the following statements is either true or false.
Delete the current content of the cell below the statement.
Then type in your answer as either `True` or `False`
(with a captial `T` or `F`).

* A counted loop is designed to iterate a specific number of times.

<!-- grade_id: cell-03f2c52fc875fd45, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-dcf0aeae52ba989d, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* Information that  is stored  and manipulated  on a  computer is called data.

<!-- grade_id: cell-4021f8b7ab820cc0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-47374c512dfb59e9, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* Since floating point numbers are extremely accurate, they should generally be used instead of `int`s.

<!-- grade_id: cell-f1f5d169efc4ba42, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-2158cbbae553bb8e, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* Operations like  addition and  subtraction are  defined  in the `math` library.

<!-- grade_id: cell-6baa652254f2b192, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-1771941c74e30f5d, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* The `sqrt()` function computes the squirt of a number.

<!-- grade_id: cell-b946c84d4f5090e4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-0774378a100ebc64, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* The `int` data type is identical to the mathematical concept of integer.

<!-- grade_id: cell-8fcd1146f01a9604, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-1ce10653dfb71f61, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* In `Python` 3, the value of `4/5` is the same as that of  `4.0/5.0`.

<!-- grade_id: cell-ce4589ef8ab553ed, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-3bc55bc442556b1b, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```