### CS102/CS103

Prof. Götz Pfeiffer<br />
School of Mathematics, Statistics and Applied Mathematics<br />
NUI Galway

# Computer Lab 3

## Multiple Choice

For each of the following questions, choose `'a'`, `'b'`, `'c'`, or `'d'` as answer
and type it into the cell below (with quotes).

* The template 
```
for <variable> in range( <expr> ):
    <body>
```
describes a

(a) a general loop<br />
(b) an assignemnent statement<br />
(c) a flowchart<br />
(d) a counted loop<br />

<!-- grade_id: cell-96a7d6eca134cef2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-fc977ed8c2f36d06, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```

* Which of the following is not a built-in operation in `python`?

(a) `+`<br />
(b) `%`<br />
(c) `sqrt()`<br />
(d) `abs()`<br />

<!-- grade_id: cell-2494c251693a75d8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-8134c2e463c21f04, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```

* In order to use functions from the `math` library, a program must include

(a) a comment<br />
(b) a loop<br />
(c) an operator<br />
(d) an `import` statement<br />

<!-- grade_id: cell-8fbaed3d68a4720c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-f6722a93060a273f, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```