### CS102/CS103

Prof. Götz Pfeiffer<br />
School of Mathematics, Statistics and Applied Mathematics<br />
NUI Galway

# Computer Lab 3

##  Programming Exercises

In this part of the practical you will write some simple `python` functions
that use a `for` loop to compute a certain value.
In each function definition, the first line of
the function body should be a string that briefly explains the
purpose of that function.

In some of the exercises, a **partial answer** together with the text
```
# YOUR CODE HERE
raise NotImplementedError()
```
is already given in the answer cell.  **Delete** those two lines
and replace them with your own code to complete the answer.

* Using the **accumulator pattern**, write a `python` function that computes the sum of the squares
of the numbers from $1$ to $n$.

<!-- grade_id: cell-2faf0e81992f1428, points: 0, grade: false, solution: true -->
```python
def sum_of_squares(n):
    "Compute the sum of the squares of the numbers from 1 to n"
    total = 0
    # YOUR CODE HERE
    raise NotImplementedError()
```

The sum of squares from $1$ to $10$ is $385$. Verify that this is the answer you get:

```python
sum_of_squares(12)
```

<!-- grade_id: cell-ef03fd263fd71770, points: 2, grade: true, solution: false -->
```python
"""Check that sum_of_squares returns the correct answer for various inputs."""
assert sum_of_squares(0) == 0
assert sum_of_squares(1) == 1
assert sum_of_squares(2) == 5
assert sum_of_squares(10) == 385
assert sum_of_squares(11) == 506

assert sum_of_squares.__doc__ != None
```

* Using the **accumulator pattern**, write a `python` function that computes the sum of the **cubes**
of the numbers from $1$ to $n$.

<!-- grade_id: cell-240ebe51f6830eb4, points: 0, grade: false, solution: true -->
```python
def sum_of_cubes(n):
    "Compute the sum of the cubes of the numbers from 1 to n"
    total = 0
    # YOUR CODE HERE
    raise NotImplementedError()
```

The sum of cubes from $1$ to $10$ is $3025$. Verify that this is the answer you get:

```python
sum_of_cubes(10)
```

<!-- grade_id: cell-b13af065382449ff, points: 2, grade: true, solution: false -->
```python
"""Check that sum_of_cubes returns the correct answer for various inputs."""
assert sum_of_cubes(0) == 0
assert sum_of_cubes(1) == 1
assert sum_of_cubes(2) == 9
assert sum_of_cubes(10) == 3025
assert sum_of_cubes(11) == 4356

assert sum_of_cubes.__doc__ != None
```

* Write a `python` function `sum_of_numbers` that computes the sum of a given list of numbers 

<!-- grade_id: cell-2d8b4080be6ecf27, points: 0, grade: false, solution: true -->
```python
def sum_of_numbers(numbers):
    "compute the sum of the list numbers"
    total = 0
    # YOUR CODE HERE
    raise NotImplementedError()
```

The sum of the list `[1, 7, 2]` is $10$. Verify that this is the answer you get:

```python
sum_of_numbers([1, 7, 2])
```

<!-- grade_id: cell-5fa27de2aeb00e13, points: 2, grade: true, solution: false -->
```python
"""Check that sum_of_numbers returns the correct answer for various inputs."""
assert sum_of_numbers([]) == 0
assert sum_of_numbers([1]) == 1
assert sum_of_numbers([1,2]) == 3
assert sum_of_numbers([1, 7, 2]) == 10
assert sum_of_numbers(range(11)) == 55
assert sum_of_numbers([i*i for i in range(11)]) == 385

assert sum_of_numbers.__doc__ != None
```

* Write a `python` function `avg_of_numbers` that computes the average of a given list of numbers. 

<!-- grade_id: cell-e0be5549829534d8, points: 0, grade: false, solution: true -->
```python
def avg_of_numbers(numbers):
    "compute the average of the list numbers"
    total = 0
    # YOUR CODE HERE
    raise NotImplementedError()
```

The average of the list `[1, 7, 2, 6]` is $4$. Verify that this is the answer you get:

```python
avg_of_numbers([1, 7, 2, 6])
```

<!-- grade_id: cell-7cc6f78c8e2ab6b2, points: 2, grade: true, solution: false -->
```python
"""Check that sum_of_numbers returns the correct answer for various inputs."""
assert avg_of_numbers([1]) == 1
assert avg_of_numbers([1,2]) == 3 / 2
assert avg_of_numbers([1, 7, 2]) == sum([1, 7, 2]) / 3
assert avg_of_numbers(range(11)) == sum(range(11)) / 11
assert avg_of_numbers([i*i for i in range(11)]) == sum([i*i for i in range(11)]) / 11

assert avg_of_numbers.__doc__ != None
```

* The mathematical constant $\pi$ can be approximated by summing the terms of the [Madhava series](https://en.wikipedia.org/wiki/Madhava_series)
$$ \frac41 - \frac43 + \frac45 - \frac47 + \frac49 - \frac4{11} \pm \dotsm $$
Write a `python` function `pi_approximation` that computes an approximation of $\pi$
as the sum of the first $n$ terms of that series.

<!-- grade_id: cell-17bde505efc70b3e, points: 0, grade: false, solution: true -->
```python
def pi_approximation(n):
    "approximate pi"
    # YOUR CODE HERE
    raise NotImplementedError()
```

The approximation of $\pi$ obtained from the first two terms of the series is $\tfrac83$. Verify that this is the answer you get:

```python
pi_approximation(2)
```

<!-- grade_id: cell-64c550d19f1557ee, points: 2, grade: true, solution: false -->
```python
"""Check that pi_approximation returns the correct answer for various inputs."""
assert pi_approximation(0) == 0.0
assert pi_approximation(1) == 4.0
assert pi_approximation(2) == 4 - 4/3
assert pi_approximation(10) == 4 * sum((-1)**i / (2*i+1) for i in range(10))

assert pi_approximation.__doc__ != None
```

* Write a `python` function that computes the factorial $n!$ of an integer $n \geq 0$
as the product of all numbers from $1$ to $n$.

<!-- grade_id: cell-88503156a3c43f7d, points: 0, grade: false, solution: true -->
```python
def factorial(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

The factorial $5!$ is $120$.  Verify that this is the answer you get:

```python
factorial(5)
```

<!-- grade_id: cell-94be2dda2f6911f0, points: 2, grade: true, solution: false -->
```python
"""Check that factorial returns the correct answer for various inputs."""
assert factorial(0) == 1
assert factorial(1) == 1
assert factorial(2) == 2
assert factorial(5) == 2 * 3 * 4 * 5

assert factorial.__doc__ != None
```

* The mathematical constant $e$ can be approximated by summing the terms of the series
$$\frac1{0!} + \frac1{1!} + \frac1{2!} + \frac1{3!} + \cdots$$
Write a `python` function that uses the above `factorial()` function and computes an approximation of $e$ as the sum of the first $n$
terms of this series.

<!-- grade_id: cell-674e36c5110c36a2, points: 0, grade: false, solution: true -->
```python
def e_approximation(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

The approximation obtained as sum of the first $3$ terms of the series is $2.5$. Verify that this is the answer you get:

```python
e_approximation(3)
```

<!-- grade_id: cell-58494e99b198abf8, points: 2, grade: true, solution: false -->
```python
"""Check that e_approximation returns the correct answer for various inputs."""
assert e_approximation(0) == 0.0
assert e_approximation(1) == 1.0
assert e_approximation(2) == 2.0
assert e_approximation(3) == 5 / 2
assert e_approximation(10) == sum(1 / factorial(i) for i in range(10))
assert e_approximation(100) == sum(1 / factorial(i) for i in range(100))

assert e_approximation.__doc__ != None
```

The [Fibonacci numbers](https://en.wikipedia.org/wiki/Fibonacci_number) are a sequence of numbers $F_n$,
$n = 0, 1, 2,  \ldots$, defined by the recurrence relation
$$F_n = F_{n-1} + F_{n-2}$$
with seed values $F_0 = 1$ and $F_1 = 1$.
Using the accumulator pattern with a list `fibonacci` initialized
as
```python
fibonacci = [0, 1]
```
Write a `python` function `fibonacci_up_to` that computes and returns
a list of Fibonacci numbers $F_0, \dots, F_n$ for a given integer $n > 0$.

<!-- grade_id: cell-d45cf919b36be9d9, points: 0, grade: false, solution: true -->
```python
def fibonacci_up_to(n):
    "compute the list of Fibonacci numbers F_0,...,F_n"
    fibonacci = [0, 1]
    # YOUR CODE HERE
    raise NotImplementedError()
```

The Fibonacci series $F_n$ up to $n = 5$ is $0, 1, 1, 2, 3, 5$. Verify that this is the answer you get:

```python
fibonacci_up_to(5)
```

<!-- grade_id: cell-343b193187dfa7e7, points: 2, grade: true, solution: false -->
```python
"""Check that fibonacci returns the correct answer for various inputs."""
assert fibonacci_up_to(1) == [0, 1]
assert fibonacci_up_to(2) == [0, 1, 1]
assert fibonacci_up_to(3) == [0, 1, 1, 2]
assert fibonacci_up_to(5) == [0, 1, 1, 2, 3, 5]
assert fibonacci_up_to(20)[20] == 6765

assert fibonacci_up_to.__doc__ != None
```

```python

```