### CS102/CS103

Prof. Götz Pfeiffer<br />
School of Mathematics, Statistics and Applied Mathematics<br />
NUI Galway

# Computer Lab 6

## True or False

Each of the following statements is either true or false.
Delete the current content of the cell below the statement.
Then type in your answer as either `True` or `False`
(with a capital `T` or `F`).

* A `python` string literal is always enclosed in double quotes

<!-- grade_id: cell-a47f83dcd3a6b7b7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-149ce7b37c6a3e3e, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* The last character of a string `s` is at position `len(s) - 1`.

<!-- grade_id: cell-a47f83dcd3a6b7b6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-249ce7b37c6a3e3e, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* A string always contains a single line of text.

<!-- grade_id: cell-a47f83dcd3a6b7b5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-349ce7b37c6a3e3e, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* In `python`, "4" + "5" is "45".

<!-- grade_id: cell-a47f83dcd3a6b7b4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-449ce7b37c6a3e3e, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* Unicode is a standard for repesenting characters using numeric codes.

<!-- grade_id: cell-a47f83dcd3a6b7b3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-549ce7b37c6a3e3e, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* The `split()` method breaks a string into a list of substrings, and `join()` does the opposite.

<!-- grade_id: cell-a47f83dcd3a6b7b2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-049ce7b37c6a3e3e, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* The `add()` method can be used to add an item to the end of a list.

<!-- grade_id: cell-1b8e9a64750602a1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-0ddf307f9dab9181, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

```python

```