### CS102/CS103

Prof. Götz Pfeiffer<br />
School of Mathematics, Statistics and Applied Mathematics<br />
NUI Galway

# Computer Lab 6

## Multiple Choice

For each of the following questions, choose `'a'`, `'b'`, `'c'`, or `'d'` as answer
and type it into the cell below (with quotes).

* Accessing a single character out of a string is called

(a) slicing<br />
(b) concatenation<br />
(c) assignment<br />
(d) indexing<br />

<!-- grade_id: cell-e2d7f46305b54d6a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-07cfc245d8db90a1, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```

* If `s` is a string, which of the following is the same as `s[0:-1]`?

(a) `s[-1]`<br />
(b) `s[:]`<br />
(c) `s[:len(s)-1]`<br />
(d) `s[0:len(s)]`<br />

<!-- grade_id: cell-e2c7f46305b54d6a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-07cfc245d8db90a2, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```

* What function gives the Unicode value of a character?

(a) `ord()`<br />
(b) `ascii()`<br />
(c) `chr()`<br />
(d) `eval()`<br />

<!-- grade_id: cell-e2b7f46305b54d6a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-07cfc245d8db90a3, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```

* The string slots that are filled in by the `format()` method are marked by:

(a) `%`<br />
(b) `$`<br />
(c) `[]`<br />
(d) `{}`<br />

<!-- grade_id: cell-e2a7f46305b54d6a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-07cfc245d8db90a4, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```

* Which of the following is **not** a file reading method in `python`?

(a) `read()`<br />
(b) `readline()`<br />
(c) `readall()`<br />
(d) `readlines()`<br />

<!-- grade_id: cell-6c2d800ce6bfc0ec, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-0c2cffbe60f28e9d, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```