### CS102/CS103

Prof. Götz Pfeiffer<br />
School of Mathematics, Statistics and Applied Mathematics<br />
NUI Galway

# Computer Lab 6

##  Programming Exercises

In this part of the practical you will write some `python` functions.

In each function definition, the **first line of
the function body** should be **a string that briefly explains the
purpose** of that function.




In some of the exercises, a **partial answer** together with the text
```
# YOUR CODE HERE
raise NotImplementedError()
```
is already given in the answer cell.  **Delete** those two lines
and replace them with your own code to complete the answer.

* An **acronym** is a word formed by taking the first letters in a phrase.
For example, RAM is an acronym for "Random Access Memory".
Write a `python` function that computes and returns an acronym for any given phrase.
The acronym should be all uppercase, even if the words in the phrase are not capitalized.
(**Hint:** use an accumulator.)

<!-- grade_id: cell-e73860bd998415ef, points: 0, grade: false, solution: true -->
```python
def acronym(phrase):
    "Compute an uppercase acronym for a phrase"
    # YOUR CODE HERE
    raise NotImplementedError()
```

So the acronym for "random access memory" is RAM. Verify that this is what your function computes.

```python
acronym("random access memory")
```

<!-- grade_id: cell-5d7c54c479512d00, points: 2, grade: true, solution: false -->
```python
"""Check that acronym returns the correct answer for various inputs"""
assert acronym("") == ""
assert acronym("a b c") == "ABC"
assert acronym("random access memory") == "RAM"
assert acronym("Random Access Memory") == "RAM"


"""Check that the function has a docstring"""
assert acronym.__doc__ != None
```

A word is an **anagram** of another word if it is formed by rearranging the letters of that word, using all the original letters exactly once.  Write a `python` function `is_anagram(word, anagram)` that returns `True` or `False`
depending on whether the argument `anagram` is an anagram of the argument `word`. In this sense,
`"silent"` is an anagram of `"listen"`.  Your function should treat uppercase letters as equal to their
lowercase versions.

<!-- grade_id: cell-a60b1eaa7c1029a5, points: 0, grade: false, solution: true -->
```python
def is_anagram(word, anagram):
    "Is 'anagram' an anagram of 'word'?"
    # YOUR CODE HERE
    raise NotImplementedError()
```

So `"Stressed"` is an anagram of `"Desserts"`. Verify that this is what your function computes.

```python
is_anagram("Desserts", "Stressed")
```

<!-- grade_id: cell-41bb02803778c6da, points: 2, grade: true, solution: false -->
```python
"""Check that is_anagram returns the correct answer for various inputs"""
assert is_anagram("", "") == True
assert is_anagram("puppet", "muppet") == False
assert is_anagram("listen", "silent") == True
assert is_anagram("Stressed", "Desserts") == True

"""Check that the function has a docstring"""
assert is_anagram.__doc__ != None
```

A word is a **palindrome** if it reads the same forwards and backwards.  Write a `python` function `is_palindrome(word)`
that checks whether a given word is a palindrome and returns `True` or `False` accordingly. Your function should treat uppercase letters as equal to their lowercase versions.
(**Hint.** The 3-argument version of slicing, `l[start:stop:step]` takes an additional argument `step` which instructs
`python` to extract values from a list `l` with step-size `step`. Thus `l[::2]` will extract all
items in `l` at even positions.  Perhaps a negative step-size will reverse a string.)

<!-- grade_id: cell-0f0c8781f44cf5c2, points: 0, grade: false, solution: true -->
```python
def is_palindrome(word):
    # YOUR CODE HERE
    raise NotImplementedError()
```

For example, `"Anna"` is a palindromic name.  Verify that this is what your function computes.

```python
is_palindrome("Anna")
```

<!-- grade_id: cell-cc5840e2854fdd36, points: 2, grade: true, solution: false -->
```python
"""Check that is_palindrome returns the correct answer for various inputs"""
assert is_palindrome("") == True
assert is_palindrome("Anna") == True
assert is_palindrome("otto") == True
assert is_palindrome("I") == True
assert is_palindrome("racecar") == True
assert is_palindrome("abc") == False

"""Check that the function has a docstring"""
assert is_palindrome.__doc__ != None
```

A **prime number** is an integer bigger than $1$ which is not divisible
by any positive integer except for $1$ and itself.  Write a `python` function `is_prime(n)` that
uses a **definite loop** to implement this definition and
returns `True` or `False`, depending on whether the given integer `n` is a prime number or not.

<!-- grade_id: cell-9f6b14955decce32, points: 0, grade: false, solution: true -->
```python
def is_prime(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

According to this definition, $1$ is not a prime, but $2$ is one. Verify that this is what your function computes.

```python
is_prime(1), is_prime(2)
```

<!-- grade_id: cell-27bd55407531fcb8, points: 2, grade: true, solution: false -->
```python
"""Check that is_prime returns the correct answer for various inputs"""
assert is_prime(1) == False
assert is_prime(2) == True
assert is_prime(3) == True
assert is_prime(4) == False
assert is_prime(5) == True
assert is_prime(6) == False
assert is_prime(7) == True
assert is_prime(0) == False
assert is_prime(-1) == False
assert is_prime(1031) == True
assert is_prime(1037) == False

"""Check that the function has a docstring"""
assert is_prime.__doc__ != None
```

Write a Python function `primes_upto(n)` that
takes an integer $n$ as its only argument
and uses your function `is_prime()` from the previous task to
compute and return the  list of all
prime numbers between $1$ and $n$ (inclusive).
(**Hint:** use an accumulator.)

<!-- grade_id: cell-6b3da3388be830d0, points: 0, grade: false, solution: true -->
```python
def primes_upto(n):
    "list primes up to n (inclusive)"
    # YOUR CODE HERE
    raise NotImplementedError()
```

The list of primes up to $n = 11$ (inclusive) is $2, 3, 5, 7, 11$. Verify that this is what your function computes.

```python
primes_upto(11)
```

<!-- grade_id: cell-10e95947a4679c74, points: 2, grade: true, solution: false -->
```python
"""Check that primes_upto returns the correct answer for various inputs"""
assert primes_upto(1) == []
assert primes_upto(2) == [2]
assert primes_upto(11) == [2, 3, 5, 7, 11]
assert len(primes_upto(1000)) == 168

"""Check that the function has a docstring"""
assert primes_upto.__doc__ != None
```

 A year is a **leap year** if it is divisible by $4$, unless it is divisible by $100$ but not by $400$.
(For example, $1800$ and $1900$ were not leap years, whereas $2000$ was one.)
Write a Python function `is_leap_year(year)` that returns
`True` or `False`, depending on whether its argument
`year` is a leap year or not.

<!-- grade_id: cell-6119b67bced5d821, points: 0, grade: false, solution: true -->
```python
def is_leap_year(year):
    # YOUR CODE HERE
    raise NotImplementedError()
```

So $2000$ was a leap year, but $2018$ won't be one.  Verify that this is what your function computes.

```python
is_leap_year(2000), is_leap_year(2018)
```

<!-- grade_id: cell-4a0b587b3f209145, points: 2, grade: true, solution: false -->
```python
"""Check that is_leap_year returns the correct answer for various inputs"""
assert is_leap_year(1900) == False
assert is_leap_year(1800) == False
assert is_leap_year(2000) == True
assert is_leap_year(2100) == False
assert is_leap_year(2016) == True
assert is_leap_year(2017) == False
assert is_leap_year(2015) == False
assert is_leap_year(2018) == False
assert is_leap_year(2019) == False

"""Check that the function has a docstring"""
assert is_leap_year.__doc__ != None
```

Write a `python` function `is_triangular(a, b, c)`
that returns `True` or `False` depending on whether
the three given numbers
`a`,
`b` and
`c` could be the lengths of the sides
of a triangle, i.e., if the sum of the smaller two
is greater than the largest of them.

<!-- grade_id: cell-9f8ab7a0659742b9, points: 0, grade: false, solution: true -->
```python
def is_triangular(a, b, c):
    # YOUR CODE HERE
    raise NotImplementedError()
```

You've probably seen a triangle with side lengths $3$, $4$, and $5$.  That makes the triple $3, 4, 5$ triangular. Verify that this is what your function computes.

```python
is_triangular(3, 4, 5)
```

<!-- grade_id: cell-9e886e3db2d7e61d, points: 2, grade: true, solution: false -->
```python
"""Check that is_triangular returns the correct answer for various inputs"""
assert is_triangular(3,4,5) == True
assert is_triangular(3,4,6) == True
assert is_triangular(3,4,7) == False
assert is_triangular(3,4,8) == False

"""Check that the function has a docstring"""
assert is_triangular.__doc__ != None
```

```python

```