### CS102/CS103

Prof. Götz Pfeiffer<br />
School of Mathematics, Statistics and Applied Mathematics<br />
NUI Galway

# Computer Lab 7

## True or False

Each of the following statements is either true or false.
Delete the current content of the cell below the statement.
Then type in your answer as either `True` or `False`
(with a capital `T` or `F`).

* The process of associating a file with an object in a program is called
"reading" the file.

<!-- grade_id: cell-4c443cd071198743, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-5807e01d05fd2070, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* Programmers rarely define their own functions.

<!-- grade_id: cell-0145d174c3a0cbfd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-b22a1fdede245744, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* A function may only be called once in a program.

<!-- grade_id: cell-e0df4636fd6a7636, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-6a355d9bdedee598, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* Information can be passed into a function through parameters.

<!-- grade_id: cell-4f889c6da1615b47, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-fa36a89cba9b8d1b, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* Every function in `python` returns some value.

<!-- grade_id: cell-fe1d02ac243eb1f4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-2bdb1634a9f5eb86, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* `python` functions can never modify a parameter.

<!-- grade_id: cell-60d6f5a284174248, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-87f56f27e30540a9, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

* One reason to use functions is to reduce code duplication.

<!-- grade_id: cell-7a033de645e3e45b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not
modify the content of that cell.)

<!-- grade_id: cell-fc43cefea4b95148, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == bool)
```

```python

```