### CS102/CS103

Prof. Götz Pfeiffer<br />
School of Mathematics, Statistics and Applied Mathematics<br />
NUI Galway

# Computer Lab 7

## Multiple Choice

For each of the following questions, choose `'a'`, `'b'`, `'c'`, or `'d'` as answer
and type it into the cell below (with quotes).

* Before reading from or writing to a file, a file object must be created via

(a) `open()`<br />
(b) `create()`<br />
(c) `File()`<br />
(d) `Folder()`<br />

<!-- grade_id: cell-fc236e455cec90ab, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-e9250e3a28f889bc, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```

* The part of a program that uses a function is called the

(a) user<br />
(b) caller<br />
(c) callee<br />
(d) statement<br />

<!-- grade_id: cell-72ecc7a84b262a9d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-c1f83bc1a99b73a8, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```

* A `python` function definition begins with the keyword

(a) `def`<br />
(b) `define`<br />
(c) `function`<br />
(d) `defun`<br />

<!-- grade_id: cell-7d115990d9e60d58, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-786091b9e291cb59, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```

* Which of the following is **not** a step in the function-calling process?

(a) The calling statement suspends execution.<br />
(b) The actual arguments are assigned to the formal parameters.<br />
(c) The body of the function is executed.<br />
(d) Control returns to the point just before the function was called.<br />

<!-- grade_id: cell-b10537fa4bca8054, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-952c0eec557e9fb2, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```

* A function with no `return` statement returns

(a) nothing<br />
(b) its parameters<br />
(c) its variables<br />
(d) `None`<br />

<!-- grade_id: cell-009735a29e21cc97, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(You can use the cell below to ensure that your answer above has the correct format. You can and should not modify the content of that cell.)

<!-- grade_id: cell-dc5c431a5290e425, points: 1, grade: true, solution: false -->
```python
_answer = _
assert(type(_answer) == str)
```

```python

```