### CS102/CS103

Prof. Götz Pfeiffer<br />
School of Mathematics, Statistics and Applied Mathematics<br />
NUI Galway

# Computer Lab 7

##  Programming Exercises

In this part of the practical you will write some `python` functions.

In each function definition, the **first line of
the function body** should be **a string that briefly explains the
purpose** of that function.



In some of the exercises, a **partial answer** together with the text
```
# YOUR CODE HERE
raise NotImplementedError()
```
is already given in the answer cell.  **Delete** those two lines
and replace them with your own code to complete the answer.

* (3 points) [Heron's formula](https://en.wikipedia.org/wiki/Heron%27s_formula) 
states that the area of a triangle whose sides have lengths $a$, $b$, and $c$ is
$$A={\sqrt {s(s-a)(s-b)(s-c)}},$$
where $s$ is the semiperimeter of the triangle; that is,
$$
s = \frac{a + b + c}{2}.
$$
Write a `python` function `area_triangle` that computes and returns the area of a triangle
from its three side lengths.  
Use last week's function `is_triangular` to check whether te given side lengths do actually
belong to a triangle, and to return `None` if not (rather than producing an error message).

<!-- grade_id: cell-206d1a6af4eecb3d, points: 0, grade: false, solution: true -->
```python
from math import sqrt
def area_triangle(a, b, c):
    "Compute area of triangle with side lengths a, b, c"
    # YOUR CODE HERE
    raise NotImplementedError()
```

The area of the triangle with side lengths $3$, $4$, and $5$ is $36$. Verify that this is what your function computes.
Also test your function with side lengths that do not form a triangle.

```python
area_triangle(1,1,2)
```

<!-- grade_id: cell-f113ef520638f83d, points: 3, grade: true, solution: false -->
```python
"""Check that area_triangle returns the correct answer for various inputs"""
assert area_triangle(3, 4, 5) == 36.0
assert area_triangle(4, 4, 4) == 48.0
assert area_triangle(5, 12, 13) == 900.0
assert area_triangle(1, 1, 2) == None

"""Check that the function has a docstring"""
assert area_triangle.__doc__ != None
```

* (3 points) The [Body Mass Index](https://en.wikipedia.org/wiki/Body_mass_index) (BMI) is calculated as a person's weight
(in pounds) times $703$, divided by the square of the person's height (in inches).
A BMI in the range from $18.5$ to $25$ (inclusive) is considered **normal**.
A BMI below $18.5$ is considered **underweight**, above $25$ and up to $30$ (inclusive) is **overweight**
and above $30$ is **obese**.  Write a function `bmi` that calculates a person's BMI from their
weight (in pounds) and their height (in inches) and returns a `python` **dictionary**
with keys `"index"` and `"result"`.  The value of the `"index"` key should be the numerical value
of the computed BMI, and the value of the `"result"` key should be one of the strings
`"normal"`, `"underweight"`, `"overweight"`, or `"obese"`, as appropriate.
Note that the BMI is a very simplistic measure of a person's thickness.
Some athletes have a high muscle to fat ratio and may have a BMI that is misleadingly high relative to their body fat percentage.

<!-- grade_id: cell-70560c39e9adab87, points: 0, grade: false, solution: true -->
```python
def bmi(weight, height):
    # YOUR CODE HERE
    raise NotImplementedError()
```

A person weighing 128 pounds with a height of 64 inches would have a BMI of $21.96875$ and hence count as normal.
Verify that this is what your function computes.

```python
bmi(128, 64)
```

<!-- grade_id: cell-18671b8c04867faa, points: 3, grade: true, solution: false -->
```python
"""Check that area_triangle returns the correct answer for various inputs"""
assert bmi(128, 64) == {'index': 21.96875, 'result': 'normal'}
assert bmi(18.5 * 64**2 / 19 / 37, 64)['result'] == 'normal'
assert bmi(25 * 64**2 / 19 / 37, 64)['result'] == 'normal'
assert bmi(30 * 64**2 / 19 / 37, 64)['result'] == 'overweight'
assert bmi(31 * 64**2 / 19 / 37, 64)['result'] == 'obese'
assert bmi(18 * 64**2 / 19 / 37, 64)['result'] == 'underweight'

"""Check that the function has a docstring"""
assert bmi.__doc__ != None
```

* (3 points) Write a function `is_valid_date(day, month, year)` that checks whether or not 
`day/month/year` is a valid date, and returns `True` or `False` accordingly.
For example, 31/11/2017 is not a valid date because November only has 30 days.
And 29/2/2018 is not valid because 2018 is not a leap year.
Also, any date with a negative value for `year` (or `day` or `month`) is invalid.
Use last week's function `is_leap_year` to check whether `year` is a leap year or not.

<!-- grade_id: cell-dc665cbf4486948f, points: 0, grade: false, solution: true -->
```python
def is_valid_date(day, month, year):
    # YOUR CODE HERE
    raise NotImplementedError()
```

So 6/11/2017 is a valid date. Verify that this is what your function computes.

```python
is_valid_date(6, 11, 2017)
```

<!-- grade_id: cell-0054f8ea0de4fbee, points: 3, grade: true, solution: false -->
```python
"""Check that is_valid_date returns the correct answer for various inputs"""
assert is_valid_date(6, 11, 2017)
assert is_valid_date(29, 2, 2000)
assert is_valid_date(29, 2, 2016)
assert not is_valid_date(29, 2, 1900)
assert not is_valid_date(29, 2, 2017)
assert not is_valid_date(6, 13, 2017)
assert not is_valid_date(-6, 11, 2017)
assert not is_valid_date(6, -11, 2017)
assert not is_valid_date(6, 11, -2017)

"""Check that the function has a docstring"""
assert is_valid_date.__doc__ != None
```

* (3 points) The [Collatz sequence](https://en.wikipedia.org/wiki/Collatz_conjecture)  of a natural number $n > 0$ is generated
by starting from $n$ and producing the next term in the sequence as follows: if the current number is even
the next term is half that number, and if the current number is odd, the next term is three times this number, plus one.
The process stops when the number 1 is reached.
For example, the Collatz sequence obtained for $n = 5$ is $5, 16, 8, 4, 2, 1$.
It is an open question in Mathematics, whether this sequence will always go to $1$, for any starting value $n$.
Write a function `collatz_sequence(n)` that computes and returns the Collatz sequence for a given
number `n` as a list, including `n`.  If `n` is less than $1$, the function should return `None`. 

<!-- grade_id: cell-7deed5b02fd57d0d, points: 0, grade: false, solution: true -->
```python
def collatz_sequence(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

As mentioned above, the Collatz sequence obtained for $n = 5$ is $5, 16, 8, 4, 2, 1$. Verify that this is what your function computes.

```python
collatz_sequence(5)
```

<!-- grade_id: cell-43b07980ddcac83c, points: 3, grade: true, solution: false -->
```python
"""Check that area_triangle returns the correct answer for various inputs"""
assert collatz_sequence(5) == [5, 16, 8, 4, 2, 1]
assert len(collatz_sequence(871)) == 179
assert collatz_sequence(0) == None

"""Check that the function has a docstring"""
assert collatz_sequence.__doc__ != None
```